```python

```

![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/lecture_banner.webp)
# LECTURE 01 Vector Analysis

# 514336 Engineering Electromagnetic Field

## Chapter 1 Vector Analysis
### 1.2 -1.3 เวคเตอร์หนึ่งหน่วยและระบบพิกัดฉาก
ในหัวข้อนี้ เราจะใข้ Library 'sympy' เพื่อการคำนวณแบบ symbolic โดยให้นิสิตทำการ

```python
import sympy as smp
```
อย่างไรก็ดี หากเกิด error แสดงว่าในเครื่องคอมพิวเตอร์หรือใน Colab ของนิสิตยังไม่ได้ install library ดังกล่าว ให้ install โดยใช้คำสั่ง

```python
%pip install sympy
```

```python
import sympy as smp
```

เวคเตอร์ $\mathbf{A} = 2\mathbf{a_x}-2\mathbf{a_y}+3\mathbf{a_z}=\begin{bmatrix}2\\-2\\3 \end{bmatrix}$
```python
smp.Matrix()
```

```python

```

ขนาดของ vector หรือ Norm ของเวคเตอร์ หาได้โดยใช้ method norm()

```python

```

การคูณ vector ด้วย scalar เช่น $2\mathbf{A}=2\begin{bmatrix}2\\-2\\3 \end{bmatrix}$

ส่งผลอย่างไรต่อ vector นั้น?

```python

```

**ตัวอย่าง 1.1** จงหาเวคเตอร์หนึ่งหน่วยที่ชี้จากจุดกำเนิดไปยังจุด $G(2, -2, 1)$

```python

```

การบวก/ลบเมตริกซ์ในระบบพิกัดฉาก
$$\mathbf{A}=\begin{bmatrix}2\\-2\\3\end{bmatrix},\quad
\mathbf{B}=\begin{bmatrix}1\\2\\1\end{bmatrix}
$$

พิจารณา
$\mathbf{A} + \mathbf{B}$

```python

```

$\mathbf{A} - \mathbf{B}$

**1.5 Dot Product**

$\mathbf{A}\cdot\mathbf{B}$

ใน Sympy ทำได้โดย A.dot(B)

```python

```

การ dot กับเวคเตอร์หนึ่งหน่วย เช่น $\mathbf{A}\cdot\mathbf{a_x}$

```python

```

### Activity 3
(1 คะแนน)

3.1 พิจารณาสนามเวคเตอร์ $\mathbf{G}(x, y, z)=y\mathbf{a_x}-2.5x\mathbf{a_y}+3\mathbf{a_z}$ จงเขียนฟังก์ชัน G เพื่อ return ค่าของเวคเตอร์ $\mathbf{G}$ ณ จุด $(x, y, z)$ ใด ๆ

<!-- grade_id: activity_3_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
def G(x, y, z):
    ### BEGIN SOLUTION
    return smp.Matrix([y, -2.5 * x, 3])
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ต่อไปนี้โดยเด็ดขาด
assert 4.38 <= G(1, 2, 3).norm() <= 4.39
assert 6.160 <= G(-2, 2, 4).norm() <= 6.165
```

3.2 จงเขียน code เพื่อหาองค์ประกอบเชิงสเกลาร์ของเวคเตอร์ $\mathbf{G}$ ณ จุด $\mathbf{Q}(4, 5,2)$ ในทิศทาง $\mathbf{a_N}=\frac{1}{3}(2\mathbf{a_x}+2\mathbf{a_y}-2\mathbf{a_z})$

(1 คะแนน)

<!-- grade_id: activity_3_2, points: 0, grade: false, solution: true -->
```python
def G_N():
    ### BEGIN SOLUTION
    G_Q = G(4, 5, 2)
    a_N = 1 / 3 * smp.Matrix([2, 2, -2])
    return G_Q.dot(a_N)
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ใน cell นี้โดยเด็ดขาด
assert -5.334 <= G_N() <= -5.333
```

3.3 จงหามุม $\mathbf{\theta}$ ระหว่างเวคเตอร์ $\mathbf{G}(\mathbf{r_Q})$ และ $\mathbf{a_N}$ (ตอบในหน่วยองศา)
(1 คะแนน)

Hint: ในการคำนวณค่า arccos เราสามารถใช้ฟังก์ชัน math.acos() และใช้ค่า $\pi$ จาก math.pi

<!-- grade_id: activity_3_3, points: 0, grade: false, solution: true -->
```python
import math
# ใช้ np.arccos() ในการคำนวณค่า arc cosine
def angle():
    ### BEGIN SOLUTION
    G_QdotA_N = G_N()
    G_Q = G(4, 5, 2)
    theta = math.acos(G_QdotA_N / G_Q.norm()) * 180.0 / math.pi
    return theta
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_3, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ใน cell นี้โดยเด็ดขาด
assert 117.434 <= angle() <= 117.435
```

## Chapter 2 Coulomb's Law

```python

```