![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/lecture_banner.webp)
# LECTURE 03 Coulomb's Law

# 514336 Engineering Electromagnetic Field

## กฎของคูลอมป์
**ตัวอย่าง 2.1**
กำหนดให้มีประจุ $Q=3\times 10^{-4}~\mathrm{C}$ วางอยู่ที่จุด $M(1, 2, 3)$ และประจุ $Q_2=-10^{-4}~\mathrm{C}$ วางอยู่ที่จุด $N(2, 0, 5)$ ในสุญญากาศ จงหาแรงที่กระทำกับประจุ $Q_2$ เนื่องจากประจุ $Q_1$

![](https://github.com/nayot/emag-jupyter/blob/main/pics/coulomb.png?raw=True)

**หมายเหตุ** Code ต่อไปนี้มีการเรียกใช้ค่าคงที่ทางวิทยาศาสตร์จาก package scipy หากยังไม่มี ให้ !pip install scipy

```python
import sympy as smp
from scipy.constants import epsilon_0, pi
```

```python
# กำหนด parameters ทุกตัว

```

```python
# หา Displacement Vector R_12

```

```python
# หา Unit Vector a_12

```

```python
# หาแรงจากกฎของคูลอมป์

```

```python
# (optional) หาขนาดของแรง

```

### Activity 3.1
จากโจทย์ในตัวอย่างที่ 2.1 จงหาแรงที่กระทำกับประจุ $Q_1$ เนื่องจากประจุ $Q_2$

```python
import sympy as smp
from scipy.constants import epsilon_0, pi
```

<!-- grade_id: activity_3_1, points: 0, grade: false, solution: true -->
```python
def F_1():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# ทดสอบ F_1()

```

<!-- grade_id: test_3_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
assert 29.957 <= F_1().norm() <= 29.959
```

## 2.2 ความเข้มสนามไฟฟ้า
สนามไฟฟ้า (Electric Field Intensity) เนื่องจากประจุแบบจุดที่วางอยู่ที่จุดกำเนิด

$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0}\frac{\mathbf{r}}{|\mathbf{r}|^3},\qquad\mathrm{V/m}$$

เมื่อ $\mathbf{r}$ คือ เวคเตอร์ตำแหน่งของจุดสังเกต

สนามไฟฟ้า เนื่องจากประจุแบบจุดที่วางอยู่ที่จุดใด ๆ ซึ่งระบุโดย $\mathbf{r'}$

$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0}\frac{\mathbf{r-r'}}{|\mathbf{r-r'}|^3}\qquad\mathrm{V/m}$$

**ตัวอย่าง 2.2** จงหาสนามไฟฟ้า $\mathbf{E}(x, y, z)$ เนื่องจากประจุแบบจุดขนาด $Q$ ที่วางอยู่ที่จุดกำเนิด ณ จุด $(x, y, z)$ ใด ๆ

```python
import sympy as smp
import scipy.constants as sp

# กำหนดตัวแปร "ต้น" ต่าง ๆ ที่ใช้ใน sympy
Q, x, y, z, pi, epsilon_0 = smp.symbols('Q x y z pi epsilon_0', real=True)
```

```python
# position vector r ที่กำหนดจุดสังเกต


```

```python
# สนามไฟฟ้า E จากกฎของคูลอมป์

```

```python
# ทดลองแ ทนค่า เพื่อหาค่าที่เป็นตัวเลข (Numerical Values)

```

### Activity 2.2
จงปรับแก้ Code ข้างต้น เพื่อสร้างฟังก์ชัน E() ที่จะ return สนามไฟฟ้า ณ ตำแหน่ง $(x, y, z)$ ใด ๆ เนื่องจากประจุแบบจุดที่วางอยู่ ณ ตำแหน่ง $(x', y', z')$

**หมายเหตุ** การที่กำหนดให้ต้องตอบในรูปฟังก์ชัน เพื่อประโยชน์ของการตรวจคำตอบอัตโนมัติ

```python
# เพื่อความสะดวกในการทดสอบ ยังไม่ต้องนำ Code ใส่ในฟังก์ชัน

import sympy as smp
# กำหนดตัวแปรต้นที่ต้องใช้
Q, pi, epsilon_0 = smp.symbols('Q pi epsilon_0', real=True)
x, y, z, xp, yp, zp = smp.symbols("x y z x' y' z'", real=True)

# YOUR CODE HERE

```

<!-- grade_id: activity_2_2, points: 0, grade: false, solution: true -->
```python
import sympy as smp
# เมื่อทดสอบจนเป็นที่พอใจแล้ว ให้นำ Code มาใส่ในฟังก์ชัน E() เพื่อให้ระบบสามารถตรวจได้
def E():
    Q, pi, epsilon_0 = smp.symbols('Q pi epsilon_0', real=True)
    x, y, z, xp, yp, zp = smp.symbols("x y z x' y' z'", real=True)

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Test calling E


```

<!-- grade_id: test_2_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
# Test Calling E()
from IPython.display import display # ใช้ display() จาก IPython เพื่อแสดงสมการใน notebook อย่างสวยงาม
display(E())

```

### การแทนค่าตัวเลขใน Sympy Expressions

ทำได้โดยการใช้ method .subs() และกำหนด input ของ method เป็น dict ที่ระบุค่าตัวเลขของตัวแปรที่ต้องการแทนค่า

```python
# กำหนดให้ E_sym เก็บเวคเตอร์สนามไฟฟ้า ณ จุดใด ๆ
E_sym = E()
E_sym
```

สมมติให้ประจุ $Q=10\mu\mathrm{C}$ วางอยู่ที่จุดกำเนิด ต้องการหาสนามที่ตำแหน่ง $\mathbf{r}=\begin{bmatrix}1\\2\\-1\end{bmatrix}$

```python
# แทนค่า


```

### การแปลง Sympy Expression เป็น Numerical Function

อีกวิธีหนึ่งในการคำนวณค่าตัวเลขของของสนาม $\mathbf{E(r)}$ คือ การแปลง Sympy Expression ให้เป็น Python Function ที่สามารถเรียกใช้งานได้ตามปกติ

Function ที่ใช้ในการแปลง Sympy Expression ให้เป็น Python Function คือ sympy.lambdify()

```python
import scipy.constants as sp
# กำหนด Sympy Expression E_sym โดยรับค่ามาจากฟังก์ชัน E() ก่อนหน้า
E_sym = E()

# แทนค่าตัวเลขในตัวแปรที่ทราบค่าแล้ว

```

```python
# lambdify

```

```python
# ทดสอบ Function E_func

```

### Activity 2.2.2
จงปรับปรุง Code ข้างต้น เพื่อสร้าง Python Function ที่รับค่าตำแหน่ง $(x, y, z)$ และ return ค่าของสนามไฟฟ้า $\mathbf{E(r)}$ ที่เกิดขึ้นเนื่องจากประจุแบบจุดขนาด $-2.5~\mu\mathrm{C}$ ที่วางอยู่ ณ ตำแหน่ง $(1, -1, 0)~\mathrm{m}$

```python
# ใช้ Sympy สร้าง Expression สนามไฟฟ้า E
import sympy as smp
from scipy.constants import pi, epsilon_0

# กำหนดตัวแปรต้นที่ใช้
x, y, z = smp.symbols('x y z', real=True)

# เวคเตอร์ตำแหน่ง/เวคเตอร์ขจัด


# ขนาดของประจุ

# สนามไฟฟ้า E

```

```python
# ใช้ lambdify() แปลง E เป็น Python Function

```

<!-- grade_id: test_2_2_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด

# ตรวจคำตอบ โดยการคำนวณแบบ numerical โดยตรงด้วย Numpy 
# หากนิสิตยังไม่เข้าใจส่วนนี้ ก็ยังไม่เป็นไร จะมีการกล่าวถึงต่อไปในภายหลัง

from scipy.constants import pi, epsilon_0
import numpy as np
from numpy.linalg import norm

Q = -2.5e-6
rp = np.array([1, -1, 0])

def E_numer(x, y, z):
    r = np.array([x, y, z])
    R = r - rp
    return Q / (4 * pi * epsilon_0) * R / norm(R)**3

# ตรวจคำตอบ
assert np.abs(norm(E_func(1, 1, 1)) - norm(E_numer(1, 1, 1))) <= 0.001

```

### การ Plot สนามไฟฟ้า
เพื่อความสะดวกในวิชานี้ ได้เตรียม Library "libemag" สำหรับการ Plot สนามไฟฟ้าใน 2 มิติ นิสิตสามารถ download หรือ clone library ดังกล่าวได้จาก https://github.com ดังนี้

https://github.com/nayot/libemag

ให้นิสิต download หรือ clone library "libemag" ไปไว้ใน folder/directory เดียวกันกับที่ไฟล์ ipynb นี้อยู่ หรืออาจจะสร้าง shortcut หรือ symbolic link ไปที่ไฟล์ดังกล่าว

```python
# ทดสอบว่า libemag ใช้การได้ / หากไม่มี error ถือว่าใช้ได้
from libemag.plots import field_plot
```

#### การใช้งานฟังก์ชัน field_plot()
```python
field_plot(field_func, num_grids=20,\
    x_min=-10, x_max=10, y_min=-10, y_max=10,\
    xlabel='x', ylabel='y', title='',\
    contour=False, cmap='plasma', streamline=False, normalising=True,
    normalising_factor=0.05)
```

โดยที่ field_func() เป็นฟังก์ชัน (numerical) ที่ return np.array เวคเตอร์ 2 มิติ ของสนามเวคเตอร์ที่ตำแหน่ง (x, y)

**ตัวอย่าง** มีประจุขนาด $1~\mu\mathrm{C}$ วางอยู่ที่จุดกำเนิด จงพล็อตสนามไฟฟ้า ณ ตำแหน่ง $(x, y)$ ใด ๆ ใน 2 มิติ

```python
# import libraries ที่จำเป็น
import sympy as smp
import numpy as np
from libemag.emag_initials import *
from libemag.plots import field_plot
```

```python
# กำหนดตัวแปรที่ต้องใช้ใน Sympy

```

```python
# กำหนดค่า Q, r, r'

```

```python
# กฎของคูลอมป์

```

```python
# Lambdify

```

```python
# Plot

```