```python

```

![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/lecture_banner.webp)
# LECTURE 03 Coulomb's Law

# 514336 Engineering Electromagnetic Field

## กฎของคูลอมป์
**ตัวอย่าง 2.1**
กำหนดให้มีประจุ $Q=3\times 10^{-4}~\mathrm{C}$ วางอยู่ที่จุด $M(1, 2, 3)$ และประจุ $Q_2=-10^{-4}~\mathrm{C}$ วางอยู่ที่จุด $N(2, 0, 5)$ ในสุญญากาศ จงหาแรงที่กระทำกับประจุ $Q_2$ เนื่องจากประจุ $Q_1$

![](https://github.com/nayot/emag-jupyter/blob/main/pics/coulombs.png?raw=True)

**หมายเหตุ** Code ต่อไปนี้มีการเรียกใช้ค่าคงที่ทางวิทยาศาสตร์จาก package scipy หากยังไม่มี ให้ !pip install scipy

```python
import sympy as smp
from scipy.constants import epsilon_0, pi
```

```python
# กำหนด parameters ทุกตัว

```

```python
# หา Displacement Vector R_12

```

```python
# หา Unit Vector a_12

```

```python
# หาแรงจากกฎของคูลอมป์

```

```python
# (optional) หาขนาดของแรง

```

### Activity 3.1
จากโจทย์ในตัวอย่างที่ 2.1 จงหาแรงที่กระทำกับประจุ $Q_1$ เนื่องจากประจุ $Q_2$

```python
import sympy as smp
from scipy.constants import epsilon_0, pi
```

<!-- grade_id: activity_3_1, points: 0, grade: false, solution: true -->
```python
def F_1():
    ### BEGIN SOLUTIONS
    Q1 = 3.0e-4
    Q2 = -1e-4
    r1 = smp.Matrix([1, 2, 3])
    r2 = smp.Matrix([2, 0, 5])

    R = r1 - r2
    
    return Q1 * Q2 / (4.0 * pi * epsilon_0) * R / R.norm()**3
    ### END SOLUTIONS
```

```python
# ทดสอบ F_1()

```

<!-- grade_id: test_3_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
assert 29.957 <= F_1().norm() <= 29.959
```

## 2.2 ความเข้มสนามไฟฟ้า
สนามไฟฟ้า (Electric Field Intensity) เนื่องจากประจุแบบจุดที่วางอยู่ที่จุดกำเนิด

$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0}\frac{\mathbf{r}}{|\mathbf{r}|^3},\qquad\mathrm{V/m}$$

เมื่อ $\mathbf{r}$ คือ เวคเตอร์ตำแหน่งของจุดสังเกต

สนามไฟฟ้า เนื่องจากประจุแบบจุดที่วางอยู่ที่จุดใด ๆ ซึ่งระบุโดย $\mathbf{r'}$

$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0}\frac{\mathbf{r-r'}}{|\mathbf{r-r'}|^3}\qquad\mathrm{V/m}$$

**ตัวอย่าง 2.2** จงหาสนามไฟฟ้า $\mathbf{E}(x, y, z)$ เนื่องจากประจุแบบจุดขนาด $Q$ ที่วางอยู่ที่จุดกำเนิด ณ จุด $(x, y, z)$ ใด ๆ

```python
import sympy as smp
import scipy.constants as sp

# กำหนดตัวแปร "ต้น" ต่าง ๆ ที่ใช้ใน sympy
Q, x, y, z, pi, epsilon_0 = smp.symbols('Q x y z pi epsilon_0', real=True)
```

```python
# position vector r ที่กำหนดจุดสังเกต


```

Output:
```
Matrix([
[x],
[y],
[z]])
```

```python
# สนามไฟฟ้า E จากกฎของคูลอมป์

```

```python
# ทดลองแ ทนค่า เพื่อหาค่าที่เป็นตัวเลข (Numerical Values)

```

### Activity 2.2
จงปรับแก้ Code ข้างต้น เพื่อสร้างฟังก์ชัน E() ที่จะ return สนามไฟฟ้า ณ ตำแหน่ง $(x, y, z)$ ใด ๆ เนื่องจากประจุแบบจุดที่วางอยู่ ณ ตำแหน่ง $(x', y', z')$

**หมายเหตุ** การที่กำหนดให้ต้องตอบในรูปฟังก์ชัน เพื่อประโยชน์ของการตรวจคำตอบอัตโนมัติ

```python
# เพื่อความสะดวกในการทดสอบ ยังไม่ต้องนำ Code ใส่ในฟังก์ชัน

import sympy as smp
# กำหนดตัวแปรต้นที่ต้องใช้
Q, pi, epsilon_0 = smp.symbols('Q pi epsilon_0', real=True)
x, y, z, xp, yp, zp = smp.symbols("x y z x' y' z'", real=True)

# YOUR CODE HERE

```

<!-- grade_id: activity_2_2, points: 0, grade: false, solution: true -->
```python
import sympy as smp
# เมื่อทดสอบจนเป็นที่พอใจแล้ว ให้นำ Code มาใส่ในฟังก์ชัน E() เพื่อให้ระบบสามารถตรวจได้
def E():
    Q, pi, epsilon_0 = smp.symbols('Q pi epsilon_0', real=True)
    x, y, z, xp, yp, zp = smp.symbols("x y z x' y' z'", real=True)

    ### BEGIN SOLUTIONS
    # เวคเตอร์ตำแหน่งที่จุดสังเกต
    r = smp.Matrix([x, y, z])
    
    # เวคเตอร์ตำแหน่งของจุดที่ประจุอยู่
    rp = smp.Matrix([xp, yp, zp])

    # Displacement vector
    R = r - rp

    # หา E จากกฎของคูลอมป์
    E = Q / (4 * pi * epsilon_0) * R / R.norm()**3
    return E
    ### END SOLUTIONS
```

```python
# Test calling E


```

<!-- grade_id: test_2_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
# Test Calling E()
from IPython.display import display # ใช้ display() จาก IPython เพื่อแสดงสมการใน notebook อย่างสวยงาม
display(E())

### BEGIN HIDDEN TESTS

Q, pi, epsilon_0 = smp.symbols('Q pi epsilon_0', real=True)
x, y, z, xp, yp, zp = smp.symbols("x y z x' y' z'", real=True)

### BEGIN SOLUTIONS
# เวคเตอร์ตำแหน่งที่จุดสังเกต
r = smp.Matrix([x, y, z])

# เวคเตอร์ตำแหน่งของจุดที่ประจุอยู่
rp = smp.Matrix([xp, yp, zp])

# Displacement vector
R = r - rp

# หา E จากกฎของคูลอมป์
ans = Q / (4 * pi * epsilon_0) * R / R.norm()**3

# Test

assert E() == ans

### END HIDDEN TESTS
```

Output:
```
Matrix([
[Q*(x - x')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))],
[Q*(y - y')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))],
[Q*(z - z')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))]])
```

### การแทนค่าตัวเลขใน Sympy Expressions

ทำได้โดยการใช้ method .subs() และกำหนด input ของ method เป็น dict ที่ระบุค่าตัวเลขของตัวแปรที่ต้องการแทนค่า

```python
# กำหนดให้ E_sym เก็บเวคเตอร์สนามไฟฟ้า ณ จุดใด ๆ
E_sym = E()
E_sym
```

Output:
```
Matrix([
[Q*(x - x')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))],
[Q*(y - y')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))],
[Q*(z - z')/(4*epsilon_0*pi*((x - x')**2 + (y - y')**2 + (z - z')**2)**(3/2))]])
```

สมมติให้ประจุ $Q=10\mu\mathrm{C}$ วางอยู่ที่จุดกำเนิด ต้องการหาสนามที่ตำแหน่ง $\mathbf{r}=\begin{bmatrix}1\\2\\-1\end{bmatrix}$

```python
# แทนค่า


```

### การแปลง Sympy Expression เป็น Numerical Function

อีกวิธีหนึ่งในการคำนวณค่าตัวเลขของของสนาม $\mathbf{E(r)}$ คือ การแปลง Sympy Expression ให้เป็น Python Function ที่สามารถเรียกใช้งานได้ตามปกติ

Function ที่ใช้ในการแปลง Sympy Expression ให้เป็น Python Function คือ sympy.lambdify()

```python
import scipy.constants as sp
# กำหนด Sympy Expression E_sym โดยรับค่ามาจากฟังก์ชัน E() ก่อนหน้า
E_sym = E()

# แทนค่าตัวเลขในตัวแปรที่ทราบค่าแล้ว
E_sym = E_sym.subs({
    Q:10e-6, 
    epsilon_0:sp.epsilon_0, 
    pi:sp.pi,
    xp:0, yp:0, zp:0
})
E_sym
```

Output:
```
Matrix([
[89875.5179226117*x/(x**2 + y**2 + z**2)**(3/2)],
[89875.5179226117*y/(x**2 + y**2 + z**2)**(3/2)],
[89875.5179226117*z/(x**2 + y**2 + z**2)**(3/2)]])
```

```python
# lambdify
E_func = smp.lambdify([x, y, z], E_sym)
```

```python
# ทดสอบ Function E_func
E_func(1, 2, 4)
```

Output:
```
array([[ 933.92599547],
       [1867.85199095],
       [3735.70398189]])
```

### Activity 2.2.2
จงปรับปรุง Code ข้างต้น เพื่อสร้าง Python Function ที่รับค่าตำแหน่ง $(x, y, z)$ และ return ค่าของสนามไฟฟ้า $\mathbf{E(r)}$ ที่เกิดขึ้นเนื่องจากประจุแบบจุดขนาด $-2.5~\mu\mathrm{C}$ ที่วางอยู่ ณ ตำแหน่ง $(1, -1, 0)~\mathrm{m}$

```python
# ใช้ Sympy สร้าง Expression สนามไฟฟ้า E
import sympy as smp
from scipy.constants import pi, epsilon_0

# กำหนดตัวแปรต้นที่ใช้
x, y, z = smp.symbols('x y z', real=True)

# เวคเตอร์ตำแหน่ง/เวคเตอร์ขจัด
r = smp.Matrix([x, y, z])
rp = smp.Matrix([1, -1, 0])
R = r - rp

# ขนาดของประจุ
Q = -2.5e-6

# สนามไฟฟ้า E
E = Q / (4 * pi * epsilon_0) * R / R.norm()**3
E
```

Output:
```
Matrix([
[ (22468.8794806529 - 22468.8794806529*x)/(z**2 + (x - 1)**2 + (y + 1)**2)**(3/2)],
[(-22468.8794806529*y - 22468.8794806529)/(z**2 + (x - 1)**2 + (y + 1)**2)**(3/2)],
[                     -22468.8794806529*z/(z**2 + (x - 1)**2 + (y + 1)**2)**(3/2)]])
```

```python
# ใช้ lambdify() แปลง E เป็น Python Function
E_func = smp.lambdify([x, y, z], E)
```

<!-- grade_id: test_2_2_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด

# ตรวจคำตอบ โดยการคำนวณแบบ numerical โดยตรงด้วย Numpy 
# หากนิสิตยังไม่เข้าใจส่วนนี้ ก็ยังไม่เป็นไร จะมีการกล่าวถึงต่อไปในภายหลัง

from scipy.constants import pi, epsilon_0
import numpy as np
from numpy.linalg import norm

Q = -2.5e-6
rp = np.array([1, -1, 0])

def E_numer(x, y, z):
    r = np.array([x, y, z])
    R = r - rp
    return Q / (4 * pi * epsilon_0) * R / norm(R)**3

# ตรวจคำตอบ
assert np.abs(norm(E_func(1, 1, 1)) - norm(E_numer(1, 1, 1))) <= 0.001

### BEGIN HIDDEN TESTS
assert np.abs(norm(E_func(-1, 1, 3)) - norm(E_numer(-1, 1, 3))) <= 0.001
assert np.abs(norm(E_func(1, -1, 1)) - norm(E_numer(1, -1, 1))) <= 0.001
assert np.abs(norm(E_func(10, 15, 20)) - norm(E_numer(10, 15, 20))) <= 0.001
### END HIDDEN TESTS
```

```python

```