![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/lecture_banner.webp)
# LECTURE 01 Introduction to EMAG and House Keeping

**ยินดีต้อนรับเข้าสู่รายวิชา 514336 Engineering Electromagnetic Fields**



<h2>ข้อตกลงเบื้องต้นในการเรียน</h2>

* เข้าเรียนตรงเวลา / มีระบบการเช็คชื่อ
* แนะนำให้เปิดกล้องในขณะเรียนตลอดเวลา ผู้ที่ไม่เปิดกล้องจะได้รับสิทธิเรียกตอบคำถามก่อนเสมอ
* ผู้ที่ไม่เปิดกล้อง และไม่มีส่วนร่วมในชั้นเรียน (เรียกถาม ไม่ตอบ) จะถูกเชิญออกไปที่ waiting room / ในกรณีนี้จะถูกยกเลิกการเช็คชื่อเข้าชั้นเรียน
* ติดต่อครูได้ทาง Google Classroom

<h2>อุปกรณ์การเรียน</h2>

* คอมพิวเตอร์ หรือ Tablet ที่สามารถใช้ Google Colab หรือ Jupyter Lab ได้
* (optional) สำหรับผู้ที่ต้องการใช้ Jupyter Lab (offline) แนะนำให้ติดตั้ง Anaconda (https://www.anaconda.com/download/success)
* เอกสารประกอบการเรียน / การส่งการบ้าน ให้ส่งในระบบ Google Classroom โดยทั้งหมดให้ส่งเป็น .ipynb ไฟล์ เท่านั้น
* หากยังไม่มี GitHub Account ให้สมัคร GitHub ด้วย เพื่อใช้ในการแชร์ไฟล์ภาพ โดยให้สร้าง Repository ชื่อ emag_acy67 และกำหนดสิทธิ์การเข้าถึงเป็น Public / ควรสร้าง folder ชื่อ pics ไว้ใน Repository ดังกล่าวเตรียมไว้เลย

<h2>กิจกรรม งานที่จะต้องส่ง และสัดส่วนคะแนน</h2>

* คะแนนเข้าเรียน (Class Attendance) - 10%
* In-Class Activities (ตรวจให้คะแนน) - 20%
* Homeworks - 20%
* Midterm - 25%
* Final - 25%


<h2>เกณฑ์การตัดเกรด</h2>

กด SHIFT + Enter สำหรับการ run cell ต่อไปนี้

```python
def main():
    score = input("Enter your score: ")
    score = float(score)
    print(f"Your grade: {get_grade(score)}")
    
def get_grade(score):
    """
    Returns the grade based on the input score using match-case.
    Parameters:
        score (float): The score (0-100).
    Returns:
        str: Grade (A, B+, B, C+, C, D+, D, or F).
    """
    if not 0.00 <= score <= 100.00:
        return "Invalid score. Please enter a score between 0 and 100."
    
    match score:
        case s if s >= 80.00:
            return "A"
        case s if s >= 75.00:
            return "B+"
        case s if s >= 70.00:
            return "B"
        case s if s >= 65.00:
            return "C+"
        case s if s >= 60.00:
            return "C"
        case s if s >= 55.00:
            return "D+"
        case s if s >= 50.00:
            return "D"
        case _:
            return "F"

main()
```

# บทที่ 1 การวิเคราะห์เชิงเวคเตอร์

**1.1.1 การพิมพ์ตัวแปรแทนปริมาณเชิงสเกลาร์ และประมาณเชิงเวคเตอร์**

#### Typographic Standards

| **Element**       | **Type Style**         | **Example**                     |
|--------------------|------------------------|----------------------------------|
| Functions          | Upright (Roman)       | $ \sin(x), \ln(x), f(x) $     |
| Variables          | Italic                | $ x, y, z, a, b $             |
| Constants          | Upright for special   | $ \pi, e, i $                 |
| Units & Operators  | Upright (Roman)       | $ \text{m/s}, \cdot, \nabla $ |
| Vectors            | Bold                  | $ \mathbf{v}$       |
| Matrices           | Bold Uppercase        | $ \mathbf{A}, \mathbf{B} $    |


### Activity 1
(1 คะแนน)

**คำสั่ง**
1. ให้ทุกคนเขียน "ตัวแปร (Variable) เวคเตอร์/เมตริกซ์" ด้วยตัวอักษร Roman a - z และ A - Z ด้วยมือลงในกระดาษหรือ Tablet
2. เมื่อเสร็จแล้วให้ถ่ายรูป หรือ capture เป็นไฟล์ภาพ upload เข้าไปใน GitHub Repository ของตนเอง
3. เปิดภาพใน mode 'raw' จากนั้น copy URL ไว้
4. Paste URL แทนที่คำว่า None (ระวังกระทบ syntax อื่น ๆ) เพื่อแนบไฟล์ภาพ

**ผลงาน**

![ผลงาน](None)

เมื่อเสร็จแล้วให้เพิ่ม return True ในฟังก์ชัน activity1_done()

<!-- grade_id: activity_1, points: 0, grade: false, solution: true -->
```python
def activity1_done():
    ## เพื่อทำ activity เสร็จแล้ว ให้ uncomment คำสั่ง return True ต่อไปนี้
    # return True
    
    ### BEGIN SOLUTION
    return True
    ### END SOLUTION
```

<!-- grade_id: correct_activity_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
assert activity1_done() == True
```

### Activity 2

(1 คะแนน)

การพิมพ์สมการใน markdown language สามารถทำได้โดยใช้ $\LaTeX$ syntax ที่อยู่ระหว่างเครื่องหมาย $ เช่น

```latex
$\mathbf{A}$ <-- หมายถึง math bold font (ตัวหนา) สำหรับเวคเตอร์
$\pi$ <-- หมายถึง ค่าพาย
```

**คำสั่ง**
ให้นิสิตเติมตารางต่อไปนี้ให้สมบูรณ์ จากนั้นแก้ไขฟังก์ชัน activity2_done() ให้ return True

|     ความหมายที่ต้องการ      |        รูปแบบตัวแปร      |
|-------------------------|-----------------------|
|ตัวแปร เวลา t               |                       |
|เวคเตอร์ความเร็ว v            |                       |
|เวคเตอร์ฟังก์ชันความเร็วกับเวลา v(t)|                       |
|หน่วย m/s                  |                        |
|ปริมาณสเกลาร์ อุณหภูมิเฉลี่ย T    |                        |
|ฟังก์ชัน sin(2 pi x)         |                        |

<!-- grade_id: activity_2, points: 0, grade: false, solution: true -->
```python
def activity2_done():
    ## เพื่อทำ activity เสร็จแล้ว ให้เพิ่มคำสั่ง return True
 
    ### BEGIN SOLUTION
    return True
    ### END SOLUTION
```

<!-- grade_id: correct_activity_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข cell นี้โดยเด็ดขาด
assert activity2_done() == True
```