![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/lecture_banner.webp)
# LECTURE 04 E-Field due to More Complicated Charges

# 514336 Engineering Electromagnetic Field

## ขนาดของสนามไฟฟ้าเนื่องจากประจุแบบจุด

จากสนามไฟฟ้า เนื่องจากประจุแบบจุด ณ จุดใด ๆ $\mathbf{r}$ ที่วางอยู่ที่ $\mathbf{r'}$

$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0}\frac{\mathbf{r-r'}}{|\mathbf{r-r'}|^3}\qquad\mathrm{V/m}$$

หากประจุดังกล่าววางอยู่ที่จุด origin จะได้ว่า
$$\mathbf{E(r)}=\frac{Q}{4\pi\epsilon_0 r^2}\mathbf{a}_r$$

และ
$$E(\mathbf{r})=|\mathbf{E(r)}|=\frac{Q}{4\pi\epsilon_0 r^2}\mathbf{a}_r$$


เมื่อ $\mathbf{r}$ คือ เวคเตอร์ตำแหน่งที่ชี้ไปยังจุดที่สนใจหรือจุดสังเกต $r$ คือระยะระหว่างจุดที่สนใจกับจุด origin และ $\mathbf{a}_r$ คือ Unit Vector ที่แสดงทิศทางของสนามไฟฟ้า ณ จุดที่สนใจ

**ตัวอย่าง** พล็อตขนาดของสนามไฟฟ้าเนื่องจากประจุ $Q=1~\mu\mathrm{C}$ ที่วางอยู่ที่จุดกำเนิด เทียบกับระยะทาง $r$

```python
# Numerical Function E(r)
import numpy as np
import sympy as smp
from scipy.constants import pi, epsilon_0

```

```python
# กำหนดตัวแปรต้น
x, y = smp.symbols('x y', real=True)
```

<!-- grade_id: example_1, points: 0, grade: false, solution: true -->
```python
# vector ตำแหน่ง/ขจัด
### BEGIN SOLUTIONS
r = smp.Matrix([x, y])
rp = smp.Matrix([0, 0])
R = r - rp

# ขนาดประจุ
Q = 1e-6

# สนามไฟฟ้า
E = Q / (4 * pi * epsilon_0) * R / R.norm()**3
E
### END SOLUTIONS
```

<!-- grade_id: example_2, points: 0, grade: false, solution: true -->
```python
# Lambdify
### BEGIN SOLUTIONS
Ex = smp.lambdify([x, y], E[0])
Ey = smp.lambdify([x, y], E[1])

def E_Q(x, y):
    return np.array([
        Ex(x, y), 
        Ey(x, y)
    ])
### END SOLUTIONS
```

<!-- grade_id: example_3, points: 0, grade: false, solution: true -->
```python
# field plot
### BEGIN SOLUTIONS
from libemag.plots import field_plot

field_plot(E_Q, contour=True)
### END SOLUTIONS
```

**พล็อตขนาดของสนามไฟฟ้าเทียบกับระยะทาง $x$**

ให้ $y=0$ พล็อต $|\mathbf{E}(x)|$

```python
import numpy as np
import matplotlib.pyplot as plt
import scienceplots
plt.style.use(['notebook', 'grid', 'science'])
```

<!-- grade_id: example_4, points: 0, grade: false, solution: true -->
```python
# กำหนด range ของ x
### BEGIN SOLUTIONS
N = 100
x = np.linspace(-10, 10, N)
y = np.zeros(N)

# คำนวณสนามไฟฟ้า
E_field = E_Q(x, y)

### Transpose E_field เพื่อให้ได้ (x, y)
E_T = E_field.T

# หา norm บน axis=1 (along the row)
E_N = np.linalg.norm(E_T, axis=1)
### END SOLUTIONS
```

<!-- grade_id: example_5, points: 0, grade: false, solution: true -->
```python
# Plot
### BEGIN SOLUTIONS
plt.figure(figsize=(8, 6))
plt.plot(x, E_N)
plt.ylim([0, 10e3])
plt.xlabel('$x$')
plt.ylabel('$|\mathbf{E}(x)|$')
plt.show()
### END SOLUTIONS
```

## สนามไฟฟ้าเนื่องจากประจุแบบไดโพล
### Activity 4.1
ประจุขนาด $1~\mu\mathrm{C}$ และ $-1~\mu\mathrm{C}$ วางอยู่ห่างกันเป็นระยะทาง $1~\mathrm{m}$ บนแกน $y$ ห่างจากจุด origin เท่า ๆ กัน จงปรับปรุงโปรแกรมข้างต้น เพื่อสร้างฟังก์ชันที่ return ค่าเวคเตอร์ของสนาม ณ จุดสังเกตใด ๆ (1 คะแนน)

<!-- grade_id: activity_4_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
from scipy.constants import pi, epsilon_0

def E_dipole(xv, yv):
    ### BEGIN SOLUTIONS
    x, y = smp.symbols('x y', real=True)
    r = smp.Matrix([x, y])
    rp = smp.Matrix([0, 0.5])
    rm = smp.Matrix([0, -0.5])
    Rp = r - rp
    Rm = r - rm
    E = Q / (4 * pi * epsilon_0) * Rp / Rp.norm()**3 - Q / (4 * pi * epsilon_0) * Rm / Rm.norm()**3

    # lambdify
    Ex = smp.lambdify([x, y], E[0])
    Ey = smp.lambdify([x, y], E[1])

    return np.array([Ex(xv, yv), Ey(xv, yv)])
    ### END SOLUTIONS
```

<!-- grade_id: test_4_1, points: 1, grade: true, solution: false -->
```python
# Testing - ห้ามแก้ไข Cell นี้โดยเด็ดขาด
assert -50.029 <= E_dipole(5, -4)[0] <= -50.028
assert 5.49 <= E_dipole(5, -4)[1] <= 5.50
```

### Activity 4.2
ใช้ field_plot() พล็อตสนามไฟฟ้าเนื่องจากประจุแบบไดโพลคู่นี้ (1 คะแนน)

<!-- grade_id: activity_4_2, points: 0, grade: false, solution: true -->
```python
from libemag.plots import field_plot
import matplotlib.pyplot as plt

def dipole_plot():
    ### BEGIN SOLUTIONS
    plt.figure(figsize=(8, 6))
    field_plot(E_dipole, num_grids=50, streamline=True, normalising=False, contour=True)
    
    q_pos = np.array([[0, -0.5], [0, 0.5]])
    plt.scatter(q_pos.T[0], q_pos.T[1], color=['b', 'r'], s=100)

    plt.xlim([-5, 5])
    plt.ylim([-5, 5])
    
    plt.show()
    ### END SOLUTIONS
    return True
```

<!-- grade_id: test_4_2, points: 1, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
try:
    assert dipole_plot()
except:
    raise AssertionError
```

## เปรียบเทียบสนามไฟฟ้าเนื่องจากประจุแบบจุดและประจุแบบไดโพล

### Activity 4.3
พล็อตขนาดของสนามไฟฟ้าเนื่องจากประจุแบบจุด เปรียบเทียบกับขนาดของสนามไฟฟ้าเนื่องจากประจุแบบไดโพล โดยกำหนดให้ $y=0$ (1 คะแนน)

<!-- grade_id: activity_4_3, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
import scienceplots
plt.style.use(['science', 'notebook', 'grid'])

def compare_plot():
    # กำหนด range ของ x ที่จะ plot
    N = 100
    x = np.linspace(-5, 5, N)
    y = np.zeros(N)

    # Initialise plot figure
    plt.figure(figsize=(8, 6))
    
    ### BEGIN SOLUTIONS
    # Plot สนามไฟฟ้าเนื่องจากประจุ Q
    E_PC = np.linalg.norm(E_Q(x, y).T, axis=1)
    plt.plot(x, E_PC)

    # Plot สนามไฟฟ้าเนื่องจากประจุแบบไดโพล
    E_DP = np.linalg.norm(E_dipole(x, y).T, axis=1)
    plt.plot(x, E_DP)

    # กำหนด plot scale
    plt.ylim([0, 100e3])
    plt.xlabel('$x~[\mathrm{m}]$')
    plt.ylabel('$|\mathbf{E}(x)|~\mathrm{[V/m]}$')

    return True
    ### END SOLUTIONS
```

<!-- grade_id: test_4_3, points: 1, grade: true, solution: false -->
```python
# Testing - ห้ามแก้ไข Cell นี้โดยเด็ดขาด
try:
    assert compare_plot()
except:
    raise AssertionError
```

## สนามไฟฟ้าเนื่องจากประจุแบบจุดที่วางเรียงกัน

### Activity 4.4
มีประจุขนาด $1~\mu\mathrm{C}$ จำนวน 11 ตัว วางเรียงห่างกันกันบนแนวแกน y ตัวละ $1~\mathrm{m}$ โดยสมมาตรกับแกน x จงพล็อตสนามไฟฟ้าที่ประจุดังกลางแผ่ออกมา (1 คะแนน)

```python
# ทดลองเขียน Code ที่นี่ก่อนสะดวกกว่า แล้วจึงนำไปบรรจุในฟังก์ชันคำตอบ
import sympy as smp
import numpy as np
from libemag.plots import field_plot
from libemag.emag_initials import pi, epsilon_0, norm
import matplotlib.pyplot as plt

import warnings
warnings.filterwarnings("ignore", category=RuntimeWarning)

### YOUR CODE HERE

```

<!-- grade_id: activity_4_4, points: 0, grade: false, solution: true -->
```python
import sympy as smp
import numpy as np
from libemag.plots import field_plot
from libemag.emag_initials import pi, epsilon_0, norm
import matplotlib.pyplot as plt

import warnings
warnings.filterwarnings("ignore", category=RuntimeWarning)

### BEGIN SOLUTIONS
def solution():
    def E_Q(xp, yp, x, y):
        Q = 1e-6
        r = np.array([x, y])
        rp = np.array([xp, yp])
        R = r - rp
        return Q / (4 * pi * epsilon_0) * R / norm(R)**3
    
    def E(x, y):
        yp = np.linspace(-5, 5, 11)
        E_field = np.array([0., 0.])
        for yi in yp:
            E_f = E_Q(0, yi, x, y)
            E_field += E_f
        return E_field
        
    
    field_plot(E, num_grids=25)
    
    y = np.linspace(-5, 5, 11)
    x = np.zeros(len(y))
    plt.scatter(x, y)
### END SOLUTIONS
    return True
```

<!-- grade_id: test_4_4, points: 1, grade: true, solution: false -->
```python
# Testing - ห้ามแก้ไข Cell นี้โดยเด็ดขาด

assert solution()
```