![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/problem_banner.webp)

# Problem Set for Class 04

**คำสั่ง**
1. ให้นิสิตทำโจทย์ทุกข้อใน Jupyter Lab Notebook นี้ โดยการเขียน code คำตอบ
2. การส่งงาน ให้ส่งด้วยการ upload ไฟล์ผ่านระบบ Google Classroom เท่านั้น โดยต้องกดยืนยันการส่ง "Done" ทุกครั้ง
3. การส่งงาน ให้ส่งไฟล์เดียวตาม format ที่กำหนดเท่านั้น หากนิสิตส่งมากกว่า 1 ไฟล์ ระบบจะไม่สามารถตรวจให้คะแนนได้
4. Format ชื่อไฟล์ต้องเป็นไปตามนี้เท่านั้น

**Filename Format**

<รหัสนิสิต>_problem4.ipynb

เช่น

65051234_problem4.ipynb

คำเตือน: การตั้ังชื่อไฟล์ไม่เป็นไปตาม format ดังกล่าว หรือตั้งชื่อไฟล์ผิด อาจทำให้นิสิตไม่ได้รับคะแนนสำหรับ Assignment นั้น ๆ เนื่องจากระบบจะไม่สามารถตรวจได้

## Problem 1
จงเขียนโปรแกรมเพื่อหาผลรวมของจำนวนต่อไปนี้ (2 คะแนน)

$$\sum_{m=0}^{5}\frac{1+(-1)^m}{m^2+1}$$

<!-- grade_id: problem_1, points: 0, grade: false, solution: true -->
```python
import numpy as np

def answer():
    ### BEGIN SOLUTIONS
    M = np.arange(0, 6)
    total = 0
    for m in M:
        total += (1 + (-1)**m) / (m**2 + 1)

    return total
    ### END SOLUTIONS
```

<!-- grade_id: test_1, points: 2, grade: true, solution: false -->
```python
# testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
from IPython.display import display

display(answer())

### BEGIN HIDDEN TESTS
assert 2.517 <= answer() <= 2.518 
### END HIDDEN TESTS
```

## Problem 2
**การใช้ Sympy ในการคำนวณทางคณิตศาสตร์ ด้วยวิธี Symbolic Mathematics**

**Problem 2.1** จงเขียนพจน์ (Expression) ทางคณิตศาสตร์ต่อไปนี้ โดยใช้ Sympy และเก็บไว้ใน "ตัวแปร (Variable)" f (1 คะแนน)
$$x^2 + 2x + \sqrt{\sin(x)}$$
**สังเกต**
ข้อนี้ให้เก็บพจน์ดังกล่าวไว้ในตัวแปร f ไม่ใช่ฟังก์ชัน โปรดระวัง!


<!-- grade_id: problem_2_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
### BEGIN SOLUTIONS
x = smp.symbols('x', real=True)
f = x**2 + 2 * x + smp.sqrt(smp.sin(x))
### END SOLUTIONS
```

<!-- grade_id: test_2_1, points: 1, grade: true, solution: false -->
```python
# Testing - ห้ามแก้ไข Cell นี้โดยเด็ดขาด
from IPython.display import display
display("Your answer: ", f)
### BEGIN HIDDEN TESTS
import sympy as smp
x = smp.symbols('x', real=True)
ff = x**2 + 2 * x + smp.sqrt(smp.sin(x))
assert f == ff
### END HIDDEN TESTS
```

**Problem 2.2** จงแทนค่า $x=\pi / 4$ ลงในพจน์ดังกล่าวในฟังก์ชัน ans() และ return ค่าคำตอบที่เป็นตัวเลข (Numberical value เช่น 3.1159) เท่านั้น (1 คะแนน)

<!-- grade_id: problem_2_2, points: 0, grade: false, solution: true -->
```python
def ans():
    ### BEGIN SOLUTIONS
    y = f.subs({x:smp.pi/4})
    return y.evalf()
    ### END SOLUTIONS
```

<!-- grade_id: test_2_2, points: 1, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
print(f"Your answer: {ans()}")

### BEGIN HIDDEN TESTS
from scipy.constants import pi
from numpy import sin, sqrt, abs
xv = pi / 4
correct_ans = xv**2 + 2 * xv + sqrt(sin(xv))
# print(f"Your answer: {ans()}")
print(f"Correct answer: {correct_ans}")
assert abs(ans() - correct_ans) <= 1e-6
### END HIDDEN TESTS
```

**Problem 2.3** จงหาอนุพันธ์ (Derivative) ของ f และแทนค่า $x=0.5$ จากนั้น return คำตอบที่เป็นตัวเลขเท่าน้้น (Numerical value) (2 คะแนน)

<!-- grade_id: problem_2_3, points: 0, grade: false, solution: true -->
```python
def ans():
    ### BEGIN SOLUTIONS
    x = smp.symbols('x', real=True)
    f = x**2 + 2 * x + smp.sqrt(smp.sin(x))
    dfdx = smp.diff(f, x)
    
    return dfdx.subs({x:0.5}).evalf()
    ### END SOLUTIONS
```

<!-- grade_id: ans_2_3, points: 2, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
print(f"Your answer: {ans()}")

### BEGIN HIDDEN TESTS
x = smp.symbols('x', real=True)
f = x**2 + 2 * x + smp.sqrt(smp.sin(x))
dfdx = smp.diff(f, x)
correct_ans = dfdx.subs({x:0.5}).evalf()

#print(f"Your answer: {ans()}")
print(f"Correct answer: {correct_ans}")
assert abs(ans() - correct_ans) <= 1e-6
### END HIDDEN TESTS
```

### Problem 3 
**Problem 3.1** จงใช้ Sympy ในการหาอินทิกรัลแบบไม่จำกัดเขต (Indefinite Integral) ต่อไปนี้ (2 คะแนน)
$$\int x^2(x^3+1)^{3/2}\, dx$$

<!-- grade_id: problem_3_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
def ans():
    ### BEGIN SOLUTIONS
    x = smp.symbols('x', real=True)
    f = x**2 * (x**3 + 1)**smp.Rational(3, 2)
    return smp.integrate(f, x)
    ### END SOLUTIONS
```

<!-- grade_id: test_3_1, points: 2, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell ต่อไปนี้โดยเด็ดขาด
from IPython.display import display
display("Your answer: ", ans())

### BEGIN HIDDEN TESTS
x = smp.symbols('x', real=True)
f1 = x**2 * (x**3 + 1)**(3/2)
f2 = x**2 * (x**3 + 1)**smp.Rational(3, 2)
ans1 = smp.integrate(f1, x)
ans2 = smp.integrate(f2, x)

#display("Your answer: ", ans())
display("Correct answer:", ans2)

assert (ans() == ans1) or (ans() == ans2)
### END HIDDEN TESTS
```

**Problem 3.2** จงใช้ Sympy ในการหาอินทิกรัลแบบจำกัดเขต (Definite Integral) ต่อไปนี้ (2 คะแนน)
$$\int_{-1}^{1} x^2(x^3+1)^{3/2}\, dx$$

<!-- grade_id: problem_3_2, points: 0, grade: false, solution: true -->
```python
import sympy as smp
def ans():
    ### BEGIN SOLUTIONS
    x = smp.symbols('x', real=True)
    f = x**2 * (x**3 + 1)**smp.Rational(3, 2)
    return smp.integrate(f, [x, -1, 1])
    ### END SOLUTIONS
```

<!-- grade_id: test_3_2, points: 2, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell ต่อไปนี้โดยเด็ดขาด
from IPython.display import display
display("Your answer: ", ans())

### BEGIN HIDDEN TESTS
x = smp.symbols('x', real=True)
f = x**2 * (x**3 + 1)**smp.Rational(3, 2)
ans1 = smp.integrate(f, [x, -1, 1])
ans2 = ans1.evalf()
display("Correct answer:", ans1)

assert (ans() == ans1) or (ans() == ans2)
### END HIDDEN TESTS
```

## Problem 4
**Problem 4.1** จากฟังก์ชันต่อไปนี้ จงสร้างพจน์ด้านขวามือด้วย Sympy แล้วบรรจุไว้ในตัวแปร f (1 คะแนน)
$$f(x)=\frac{2}{\sqrt{1+x^2}},\quad x=[0, 2]$$

<!-- grade_id: problem_4_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp

### BEGIN SOLUTIONS
x = smp.symbols('x', real=True)
f = 2 / smp.sqrt(1 + x**2)
### END SOLUTIONS
```

<!-- grade_id: test_4_1, points: 1, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
from IPython.display import display
display("Your answer:", f)

### BEGIN HIDDEN TESTS
import sympy as smp
x = smp.symbols('x', real=True)
f_ans = 2 / smp.sqrt(1 + x**2)
display("Correct answer:", f_ans)

assert f == f_ans
### END HIDDEN TESTS
```

**Problem 4.2** จงแปลงพจน์ข้างต้นเป็น Numerical (Numpy) function: f_func(x) โดยใช้ lambdify() method ของ Sympy (2 คะแนน)

<!-- grade_id: problem_4_2, points: 0, grade: false, solution: true -->
```python
import sympy as smp
### BEGIN SOLUTIONS
x = smp.symbols('x', real=True)
f = 2 / smp.sqrt(1 + x**2)
# Lambdify
f_func = smp.lambdify(x, f)
### END SOLUTIONS
```

<!-- grade_id: test_4_2, points: 2, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cell นี้โดยเด็ดขาด
print(f"Your answer f_func(0.5) = {f_func(0.5)}")

### BEGIN HIDDEN TESTS
x = smp.symbols('x', real=True)
f = 2 / smp.sqrt(1 + x**2)
# Lambdify
f_ans = smp.lambdify(x, f)
print(f"Correct answer f_func(0.5) = {f_ans(0.5)}")

assert f_func(0.5) == f_ans(0.5)
assert f_func(-0.5) == f_ans(-0.5)
assert f_func(15) == f_ans(15)

### END HIDDEN TESTS
```

**Problem 4.3** จงกพล็อตกราฟของฟังก์ชัน f_func(x) ในช่วง $x=[0, 2]$ (2 คะแนน)

<!-- grade_id: problem_4_3, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
import numpy as np
# อย่าลืม return True

def plot_func():
    ### BEGIN SOLUTIONS
    x = np.linspace(0, 2, 100)
    y = f_func(x)
    plt.plot(x, y)
    plt.xlabel("x")
    plt.ylabel("y")
    plt.show()

    return True
    ### END SOLUTIONS
```

<!-- grade_id: test_4_3, points: 2, grade: true, solution: false -->
```python
# Testing -- ห้ามแก้ไข Cellนี้โดยเด็ดขาด
assert plot_func()
```