![](https://raw.githubusercontent.com/nayot/emag-jupyter/refs/heads/main/pics/problem_banner.webp)

# Problem Set 1

**คำสั่ง**
1. ให้นิสิตทำโจทย์ทุกข้อใน Jupyter Lab Notebook นี้ โดยการเขียน code คำตอบ
2. การส่งงาน ให้ส่งด้วยการ upload ไฟล์ผ่านระบบ Google Classroom เท่านั้น โดยต้องกดยืนยันการส่ง "Done" ทุกครั้ง
3. การส่งงาน ให้ส่งไฟล์เดียวตาม format ที่กำหนดเท่านั้น หากนิสิตส่งมากกว่า 1 ไฟล์ ระบบจะไม่สามารถตรวจให้คะแนนได้
4. Format ชื่อไฟล์ต้องเป็นไปตามนี้เท่านั้น

**Filename Format**

<รหัสนิสิต>.ipynb

เช่น

65051234.ipynb

คำเตือน: การตั้ังชื่อไฟล์ไม่เป็นไปตาม format ดังกล่าว หรือตั้งชื่อไฟล์ผิด อาจทำให้นิสิตไม่ได้รับคะแนนสำหรับ Assignment นั้น ๆ เนื่องจากระบบจะไม่สามารถตรวจได้

### Problem 1

กำหนดจุด $M(-1, 2, 1)$ และ $N(3, -2, -2)$

#### Problem 1.1
(1 คะแนน)

จุดดังกล่าวอยู่ใน Octant ที่เท่าไร ให้นิสิต ระบุเลข Octant ใน Code ข้างล่างนี้ แทนที่คำว่า "None"

<!-- grade_id: problem_1_1, points: 0, grade: false, solution: true -->
```python
def problem_1_1():
    # ใส่คำตอบตรงนี้
    M_octant = None
    N_octant = None

    # BEGIN SOLUTION
    M_octant = 2
    N_Octant = 7
    # END SOLUTION
    return (M_octant, N_Octant)
```

<!-- grade_id: correct_problem_1_1, points: 1, grade: true, solution: false -->
```python
# เรียกฟังก์ชัน เพื่อดูคำตอบของนิสิต
problem_1_1()

### BEGIN HIDDEN TESTS
assert problem_1_1() == (2, 7)
### END HIDDEN TESTS
```

#### Problem 1.2
จากจุด $M$ และ $N$ ข้างต้น จงหาเวคเตอร์ขจัด (Displacement Vector) $\mathbf{R}_{MN}$

<!-- grade_id: problem_1_2, points: 0, grade: false, solution: true -->
```python
import sympy as smp

def problem_1_2():
    # Position Vector r_M and r_N
    r_M = None
    r_N = None
    
    # หา Displacement Vector
    R_MN = None

    #### BEGIN SOLUTION
    # Position Vector r_M and r_N
    r_M = smp.Matrix([-1, 2, 1])
    r_N = smp.Matrix([3, -2, -2])
    
    # หา Displacement Vector
    R_MN = r_N - r_M
    #### END SOLUTION
    return R_MN

```

<!-- grade_id: correct_problem_1_2, points: 1, grade: true, solution: false -->
```python
problem_1_2()

# BEGIN HIDDEN TESTS
assert problem_1_2() == smp.Matrix([4, -4, -3])
# END HIDDEN TESTS
```

#### Problem 1.3
จงเขียนฟังก์ชันคำนวณหาระยะทาง (ปริมาณสเกลาร์) ระหว่างจุดสองจุด ที่อธิบายด้วยเวคเตอร์ตำแหน่ง $\mathbf{r}_1$ และ $\mathbf{r}_2$ โดยกำหนดให้่ฟังก์ชัน รับ input และ output ในรูปแบบดังนี้

input: smp.Matrix([x, y, z]) โดยที่ x, y และ z คือตำแหน่งในระบบพิกัดฉาก

output: float

<!-- grade_id: problem_1_3, points: 0, grade: false, solution: true -->
```python
# คำนวณระยะทางระหว่างจุด M และ N
# r_M, r_N เป็น position vectors ที่อยู่ในรูป Sympy Matrix
def distance(r_M, r_N):
    ### BEGIN SOLUTION
    R = r_N - r_M
    dist = R.norm()
    ### END SOLUTION
    
    return dist
```

<!-- grade_id: correct_problem_1_3, points: 1, grade: true, solution: false -->
```python
# Test the function
import sympy as smp

r_M = smp.Matrix([0, 0, 0])
r_N = smp.Matrix([1, 1, 1])

assert distance(r_M, r_N) == smp.sqrt(3)
# ถ้า run code นี้แล้วไม่มี error แสดงว่าคำตอบถูกต้อง

### BEGIN HIDDEN TEST
r_M = smp.Matrix([-1, 5.6, 4.8])
r_N = smp.Matrix([3, -1, -1])
assert 9.650 <= distance(r_M, r_N) <= 9.655
### END HIDDEN TEST
```

### Problem 2
หากมีสนามเวคเตอร์ $\mathbf{S}$ ในระบบพิกัดฉาก ซึ่งอธิบายได้ด้วย

$$\mathbf{S}=\frac{125}{(x-1)^2+(y-2)^2+(z+1)^2} \left[(x-1)\mathbf{a_x}+(y-2)\mathbf{a_y}+(z+1)\mathbf{a_z}\right]$$

#### Problem 2.1
จงเขียน code เพื่อสร้างฟังก์ชัน S(r) ที่รับ input เป็นเวคเตอร์ตำแหน่ง และคำนวนณ output เป็นค่าของสนามเวคเตอร์ ณ ตำแหน่งนั้น

<!-- grade_id: problem_2_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
def S(r):
    # BEGIN SOLUTION
    x, y, z = r
    term1 = 125.0 / ((x - 1 )**2 + (y - 2)**2 + (z + 1)**2)
    term2 = smp.Matrix([x - 1, y - 2, z + 1])
    return term1 * term2
    # END SOLUTION
```

<!-- grade_id: correct_problem_2_1, points: 1, grade: true, solution: false -->
```python
r = smp.Matrix([2, 4, 3])
assert 27.276 <= S(r).norm() <= 27.278

### BEGIN HIDDEN TEST ###
r = smp.Matrix([-1, 0, 1])
assert 36.083 <= S(r).norm() <= 36.085
### END HIDDEN TEST ###
```

#### Problem 2.2
จงหาสร้างฟังก์ชัน เพื่อคำนวณหาเวคเตอร์หนึ่งหน่วยที่มีทิศทางเช่นเดียวกับ $\mathbf{S}$ ณ จุด $\mathbf{r_p}$ ใด ๆ

<!-- grade_id: problem_2_2, points: 0, grade: false, solution: true -->
```python
def a_S(r_p):
    ### BEGIN SOLUTION
    Sp = S(r_p)
    return Sp / Sp.norm()
    ### END SOLUTION
```

<!-- grade_id: correct_problem_2_2, points: 1, grade: true, solution: false -->
```python
import sympy as smp

r_p = smp.Matrix([2, 4, 3])
error = a_S(r_p) - smp.Matrix([0.218, 0.436, 0.873])
assert error.norm() <= 1e-3
# ถ้า run cell นี้แล้วไม่มี error แสดงว่า code ทำงานได้ถูกต้อง

### BEGIN HIDDEN TESTS
r_p = smp.Matrix([25.45, -12.03, 5.45])
error = a_S(r_p) - smp.Matrix([0.845, -0.484, 0.223])
assert error.norm() <= 1e-2
### END HIDDEN TESTS

```