![](https://raw.githubusercontent.com/nayot/emag/refs/heads/main/pics/banner.webp)
# LECTURE 01 Introduction to EMAG and House Keeping

**ยินดีต้อนรับเข้าสู่รายวิชา 514336 Engineering Electromagnetic Fields**



<h2>ข้อตกลงเบื้องต้นในการเรียน</h2>

* เข้าเรียนตรงเวลา / มีระบบการเช็คชื่อ
* แนะนำให้เปิดกล้องในขณะเรียนตลอดเวลา ผู้ที่ไม่เปิดกล้องจะได้รับสิทธิเรียกตอบคำถามก่อนเสมอ
* ผู้ที่ไม่เปิดกล้อง และไม่มีส่วนร่วมในชั้นเรียน (เรียกถาม ไม่ตอบ) จะถูกเชิญออกไปที่ waiting room / ในกรณีนี้จะถูกยกเลิกการเช็คชื่อเข้าชั้นเรียน
* ติดต่อครูได้ทาง Google Classroom

<h2>อุปกรณ์การเรียน</h2>

* คอมพิวเตอร์ หรือ Tablet ที่สามารถใช้ Google Colab หรือ Jupyter Lab ได้
* (optional) สำหรับผู้ที่ต้องการใช้ Jupyter Lab (offline) แนะนำให้ติดตั้ง Anaconda (https://www.anaconda.com/download/success)
* เอกสารประกอบการเรียน / การส่งการบ้าน ให้ส่งในระบบ Google Classroom โดยทั้งหมดให้ส่งเป็น .ipynb ไฟล์ เท่านั้น
* หากยังไม่มี GitHub Account ให้สมัคร GitHub ด้วย เพื่อใช้ในการแชร์ไฟล์ภาพ โดยให้สร้าง Repository ชื่อ emag_acy67 และกำหนดสิทธิ์การเข้าถึงเป็น Public / ควรสร้าง folder ชื่อ pics ไว้ใน Repository ดังกล่าวเตรียมไว้เลย

<h2>กิจกรรม งานที่จะต้องส่ง และสัดส่วนคะแนน</h2>

* คะแนนเข้าเรียน (Class Attendance) - 10%
* In-Class Activities (ตรวจให้คะแนน) - 20%
* Homeworks - 20%
* Midterm - 25%
* Final - 25%


<h2>เกณฑ์การตัดเกรด</h2>

กด SHIFT + Enter สำหรับการ run cell ต่อไปนี้

```python
def main():
    score = input("Enter your score: ")
    score = float(score)
    print(f"Your grade: {get_grade(score)}")
    
def get_grade(score):
    """
    Returns the grade based on the input score using match-case.
    Parameters:
        score (float): The score (0-100).
    Returns:
        str: Grade (A, B+, B, C+, C, D+, D, or F).
    """
    if not 0.00 <= score <= 100.00:
        return "Invalid score. Please enter a score between 0 and 100."
    
    match score:
        case s if s >= 80.00:
            return "A"
        case s if s >= 75.00:
            return "B+"
        case s if s >= 70.00:
            return "B"
        case s if s >= 65.00:
            return "C+"
        case s if s >= 60.00:
            return "C"
        case s if s >= 55.00:
            return "D+"
        case s if s >= 50.00:
            return "D"
        case _:
            return "F"

main()
```

# บทที่ 1 การวิเคราะห์เชิงเวคเตอร์

**1.1.1 การพิมพ์ตัวแปรแทนปริมาณเชิงสเกลาร์ และประมาณเชิงเวคเตอร์**

#### Typographic Standards

| **Element**       | **Type Style**         | **Example**                     |
|--------------------|------------------------|----------------------------------|
| Functions          | Upright (Roman)       | $ \sin(x), \ln(x), f(x) $     |
| Variables          | Italic                | $ x, y, z, a, b $             |
| Constants          | Upright for special   | $ \pi, e, i $                 |
| Units & Operators  | Upright (Roman)       | $ \text{m/s}, \cdot, \nabla $ |
| Vectors            | Bold                  | $ \mathbf{v}$       |
| Matrices           | Bold Uppercase        | $ \mathbf{A}, \mathbf{B} $    |


### Activity 1
(1 คะแนน)

**คำสั่ง**
1. ให้ทุกคนเขียน "ตัวแปร (Variable) เวคเตอร์/เมตริกซ์" ด้วยตัวอักษร Roman a - z และ A - Z ด้วยมือลงในกระดาษหรือ Tablet
2. เมื่อเสร็จแล้วให้ถ่ายรูป หรือ capture เป็นไฟล์ภาพ upload เข้าไปใน GitHub Repository ของตนเอง
3. เปิดภาพใน mode 'raw' จากนั้น copy URL ไว้
4. Paste URL แทนที่คำว่า None (ระวังกระทบ syntax อื่น ๆ) เพื่อแนบไฟล์ภาพ

**ผลงาน**

![ผลงาน](None)

เมื่อเสร็จแล้วให้เพิ่ม return True ในฟังก์ชัน activity1_done()

<!-- grade_id: activity_1, points: 0, grade: false, solution: true -->
```python
def activity1_done():
    ## เพื่อทำ activity เสร็จแล้ว ให้ uncomment คำสั่ง return True ต่อไปนี้
    # return True
    
    ### BEGIN SOLUTION
    return True
    ### END SOLUTION
```

<!-- grade_id: correct_activity_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข Cell นี้โดยเด็ดขาด
assert activity1_done() == True
```

### Activity 2

(1 คะแนน)

การพิมพ์สมการใน markdown language สามารถทำได้โดยใช้ $\LaTeX$ syntax ที่อยู่ระหว่างเครื่องหมาย $ เช่น

```latex
$\mathbf{A}$ <-- หมายถึง math bold font (ตัวหนา) สำหรับเวคเตอร์
$\pi$ <-- หมายถึง ค่าพาย
```

**คำสั่ง**
ให้นิสิตเติมตารางต่อไปนี้ให้สมบูรณ์ จากนั้นแก้ไขฟังก์ชัน activity2_done() ให้ return True

|     ความหมายที่ต้องการ      |        รูปแบบตัวแปร      |
|-------------------------|-----------------------|
|ตัวแปร เวลา t               |                       |
|เวคเตอร์ความเร็ว v            |                       |
|เวคเตอร์ฟังก์ชันความเร็วกับเวลา v(t)|                       |
|หน่วย m/s                  |                        |
|ปริมาณสเกลาร์ อุณหภูมิเฉลี่ย T    |                        |
|ฟังก์ชัน sin(2 pi x)         |                        |

<!-- grade_id: activity_2, points: 0, grade: false, solution: true -->
```python
def activity2_done():
    ## เพื่อทำ activity เสร็จแล้ว ให้เพิ่มคำสั่ง return True
 
    ### BEGIN SOLUTION
    return True
    ### END SOLUTION
```

<!-- grade_id: correct_activity_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข cell นี้โดยเด็ดขาด
assert activity2_done() == True
```

```python

```

**1.2 -1.3 เวคเตอร์หนึ่งหน่วยและระบบพิกัดฉาก**
ในหัวข้อนี้ เราจะใข้ Library 'sympy' เพื่อการคำนวณแบบ symbolic โดยให้นิสิตทำการ

```python
import sympy as smp
```
อย่างไรก็ดี หากเกิด error แสดงว่าในเครื่องคอมพิวเตอร์หรือใน Colab ของนิสิตยังไม่ได้ install library ดังกล่าว ให้ install โดยใช้คำสั่ง

```python
%pip install sympy
```

```python
import sympy as smp
```

เวคเตอร์ $\mathbf{A} = 2\mathbf{a_x}-2\mathbf{a_y}+3\mathbf{a_z}=\begin{bmatrix}2\\-2\\3 \end{bmatrix}$
```python
smp.Matrix()
```

```python

```

ขนาดของ vector หรือ Norm ของเวคเตอร์ หาได้โดยใช้ method norm()

```python

```

การคูณ vector ด้วย scalar เช่น $2\mathbf{A}=2\begin{bmatrix}2\\-2\\3 \end{bmatrix}$

ส่งผลอย่างไรต่อ vector นั้น?

```python

```

**ตัวอย่าง 1.1** จงหาเวคเตอร์หนึ่งหน่วยที่ชี้จากจุดกำเนิดไปยังจุด $G(2, -2, 1)$

```python

```

การบวก/ลบเมตริกซ์ในระบบพิกัดฉาก
$$\mathbf{A}=\begin{bmatrix}2\\-2\\3\end{bmatrix},\quad
\mathbf{B}=\begin{bmatrix}1\\2\\1\end{bmatrix}
$$

พิจารณา
$\mathbf{A} + \mathbf{B}$

```python

```

$\mathbf{A} - \mathbf{B}$

**1.5 Dot Product**

$\mathbf{A}\cdot\mathbf{B}$

ใน Sympy ทำได้โดย A.dot(B)

```python

```

การ dot กับเวคเตอร์หนึ่งหน่วย เช่น $\mathbf{A}\cdot\mathbf{a_x}$

```python

```

### Activity 3
(1 คะแนน)

3.1 พิจารณาสนามเวคเตอร์ $\mathbf{G}(x, y, z)=y\mathbf{a_x}-2.5x\mathbf{a_y}+3\mathbf{a_z}$ จงเขียนฟังก์ชัน G เพื่อ return ค่าของเวคเตอร์ $\mathbf{G}$ ณ จุด $(x, y, z)$ ใด ๆ

<!-- grade_id: activity_3_1, points: 0, grade: false, solution: true -->
```python
import sympy as smp
def G(x, y, z):
    ### BEGIN SOLUTION
    return smp.Matrix([y, -2.5 * x, 3])
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_1, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ต่อไปนี้โดยเด็ดขาด
assert 4.38 <= G(1, 2, 3).norm() <= 4.39
assert 6.160 <= G(-2, 2, 4).norm() <= 6.165
```

3.2 จงเขียน code เพื่อหาองค์ประกอบเชิงสเกลาร์ของเวคเตอร์ $\mathbf{G}$ ณ จุด $\mathbf{Q}(4, 5,2)$ ในทิศทาง $\mathbf{a_N}=\frac{1}{3}(2\mathbf{a_x}+2\mathbf{a_y}-2\mathbf{a_z})$

(1 คะแนน)

<!-- grade_id: activity_3_2, points: 0, grade: false, solution: true -->
```python
def G_N():
    ### BEGIN SOLUTION
    G_Q = G(4, 5, 2)
    a_N = 1 / 3 * smp.Matrix([2, 2, -2])
    return G_Q.dot(a_N)
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_2, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ใน cell นี้โดยเด็ดขาด
assert -5.334 <= G_N() <= -5.333
```

3.3 จงหามุม $\mathbf{\theta}$ ระหว่างเวคเตอร์ $\mathbf{G}(\mathbf{r_Q})$ และ $\mathbf{a_N}$ (ตอบในหน่วยองศา)
(1 คะแนน)

Hint: ในการคำนวณค่า arccos เราสามารถใช้ฟังก์ชัน math.acos() และใช้ค่า $\pi$ จาก math.pi

<!-- grade_id: activity_3_3, points: 0, grade: false, solution: true -->
```python
import math
# ใช้ np.arccos() ในการคำนวณค่า arc cosine
def angle():
    ### BEGIN SOLUTION
    G_QdotA_N = G_N()
    G_Q = G(4, 5, 2)
    theta = math.acos(G_QdotA_N / G_Q.norm()) * 180.0 / math.pi
    return theta
    ### END SOLUTION
```

<!-- grade_id: correct_activity_3_3, points: 1, grade: true, solution: false -->
```python
# ห้ามแก้ไข code ใน cell นี้โดยเด็ดขาด
assert 117.434 <= angle() <= 117.435
```